---
title: "Test of gradethis"
tutorial:
 id: gradethis_test
 version: 0.1
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
    css: "css/font-size.css"
runtime: shiny_prerendered
description: >
  In introduction to R for applied epidemiologists, using relevant data and examples.
  See www.appliedepi.org and www.epiRhandbook.com for more offerings from Applied Epi.
---

<!-- Add JavaScript code for making the exercise code larger -->
<script language="JavaScript" src="js/exercise-font-size.js"></script>

```{r setup, include=FALSE}
# load packages ----------------------------------------------------------------
library(learnr)
library(gradethis)
library(tidyverse)
library(here)
library(rio)
library(basket) # not sure if we need this
library(etude) # helper functions for gradethis

# set options for exercises and checking ---------------------------------------
gradethis_setup()

learnr::tutorial_options(exercise.timelimit = 60) 
  #exercise.checker = gradethis::grade_learnr) 
    # alternatively, submitr::null_code_checker


# event recorder ---------------------------------------------------------------
# see https://github.com/dtkaplan/submitr/blob/master/R/make_a_recorder.R
tutorial_options(exercise.eval = FALSE)  # pre-evaluate exercises

vfun <- submitr::make_basic_validator(NULL, "hello") #basket::check_valid

new_recorder <- function(tutorial_id, tutorial_version, user_id, event, data) {
    cat(
      tutorial_id, 
      " (v", tutorial_version, "); ",
      format(Sys.time(), "%Y-%M%-%D %H:%M:%S %Z"), "; ",
      user_id, "; ",
      event, "; ",
      data$label, "; ",
      data$answers, "; ",
      data$code, "; ",
      data$correct, "\n", sep = "",
      
      file = here::here("event_records", "test_gradethis_nb.txt"),
      append = TRUE)
}

options(tutorial.event_recorder = new_recorder)


# hide non-exercise code chunks ------------------------------------------------
knitr::opts_chunk$set(echo = FALSE)


# data prep --------------------------------------------------------------------
linelist <- rio::import(here::here("data", "linelist_cleaned.rds"))
```


```{r}
submitr::login_controls() # print login and password with "Submit" button.
```


```{r context = "server", echo = FALSE}
# see https://rdrr.io/github/dtkaplan/submitr/f/vignettes/using.Rmd
options(tutorial.storage = "none")
vfun <- submitr::make_basic_validator(NULL, "hello")       #basket::check_valid
storage_actions <- submitr::record_local("./minimal_submissions.csv")
submitr::shiny_logic(input, output, session, vfun,
                     storage_actions)
```





## Introduction to R for Applied Epidemiologists

```{r appliedepi-banner, fig.margin = TRUE, echo = FALSE, fig.width = 3, out.width = "100%", fig.cap = ""}
knitr::include_graphics("images/moz-banner.png")
```

Welcome to the course "Introduction to R for applied epidemiologists", offered for free by [Applied Epi](www.appliedepi.org) - a non-profit organisation that elevates grassroots epidemiology by offering open-source tools, training, and support to frontline public health practitioners.  

```{r appliedepi-hexes, fig.margin = TRUE, echo = FALSE, fig.width = 3, out.width = "100%", fig.cap = ""}
knitr::include_graphics("images/hex-sidebyside.png")
```

#### Other languages
This course is available...

#### Offline / Online

You can access this tutorial offline by downloading our R package ... 
If viewing offline, you can view the videos by doing ...

#### Target Audience  

This training is designed for ....
We expect that you know how to do ...

#### Learning goals

- Goal 1
- Goal 2
- Goal 3



## Packages

In this assignment we'll work with 3 R packages, let's load them!

```{r load-package, exercise = TRUE}
___(rio)
___(here)
___(tidyverse)
```

```{r load-package-solution}
library(rio)
library(here)
library(tidyverse)
```

```{r load-package-check}
grade_this_code("You are correct, the packages you need for this lesson are now loaded!")
```


## Data

Now, import the dataset "linelist_cleaned.rds" from the "data" folder.

```{r import-data, exercise = TRUE}

```

```{r import-data-solution}
linelist <- rio::import("data", "linelist_cleaned.rds")
```

```{r import-data-check}
grade_this({
  # if object is data frame
  if (is.data.frame(.result)){
    pass("linelist is a data frame")
  }
  
  # default
  fail()
})
```

### Review data

Now we can review the dataset - press "Run Code" to see the result.

```{r view-data, exercise = TRUE}
tibble(linelist)
```


```{r quiz-1}
quiz(
  question("What class is the column 'age'?",
    answer("numeric"),
    answer("double", correct = TRUE),
    answer("character"),
    answer("logical"),
    answer("proletariat"),
    incorrect = "Check the print-out under the column name, between <>",
    allow_retry = TRUE
  ),
  
question_text(
  "Please enter the function to open a dataframe in the viewier, below",
  answer("view()", message = "Remember that this function begins with a capital letter"),
  answer("print()", message = "Nope, try View() instead"),
  answer("View()", correct = TRUE),
  allow_retry = TRUE
)
)
```





## Count cases

Run a `count()` command to tabulate the number of cases per hospital, and use `arrange()` and `desc()` to order the hospitals in descending order by number of cases. 


```{r count-cases, exercise = TRUE}
linelist %>% 
  ___(___) %>% 
  ___(desc(n))
```


```{r count-cases-solution}
linelist %>% 
  count(hospital) %>% 
  arrange(desc(n))
```


```{r count-cases-check}
grade_result(
  pass_if(~ identical(as.numeric(.result[1,2]), 1762), "You successfully identified Port Hospital as the facility with the most cases!"),
  fail_if(~ identical(as.numeric(.result[1,2]), 1469), "No, Did you forget to arrange() and ensure descending order on the column 'n'?"),
  fail_if(~ TRUE)
)
```

<!-- ## Data -->

<!-- We have data from Lego sales in 2018 for a sample of customers who bought Legos in the United States. -->
<!-- The dataset is available as part of the **dsbox** package, and it's called `lego_sales`. -->

<!-- The codebook for the dataset is as follows: -->

<!-- | Name                       |  Description -->
<!-- |:----------------|:-------------------------------- -->
<!-- | `first_name` | First name of customer -->
<!-- | `last_name` | Last name of customer -->
<!-- | `age` | Age of customer -->
<!-- | `phone_number` | Phone number of customer -->
<!-- | `set_id` | Set ID of lego set purchased -->
<!-- | `number` | Item number of lego set purchased -->
<!-- | `theme` | Theme of lego set purchased -->
<!-- | `subtheme` | Sub theme of lego set purchased -->
<!-- | `year` | Year of purchase -->
<!-- | `name` |  Name of lego set purchased -->
<!-- | `pieces` | Number of pieces of legos in set purchased -->
<!-- | `us_price` | Price of set purchase in US Dollars -->
<!-- | `image_url` | Image URL of lego set purchased -->
<!-- | `quantity` | Quantity of lego set(s) purchased -->

<!-- ## Counting frequencies -->

<!-- ### Most common names -->

<!-- What are the three most common first names of purchasers? Use the code chunk below to find out! -->

<!-- ```{r common-names, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___(___) -->
<!-- ``` -->

<!-- ```{r common-names-hint-1} -->
<!-- lego_sales %>% -->
<!--   count(___, sort = TRUE) -->
<!-- ``` -->

<!-- ```{r common-names-solution} -->
<!-- lego_sales %>% -->
<!--   count(first_name, sort = TRUE) -->
<!-- ``` -->

<!-- ```{r common-names-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "Jackson")) { -->
<!--     pass("You have successfully counted names and sorted the counts in descending order.") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Aaron")) {  -->
<!--     fail("Did you forget to sort the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Adoria")) { -->
<!--     fail("Did you accidentally sort the counts in ascending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Ackah Yensu")) { -->
<!--     fail("Did you count last names instead of first names?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Brown")) {  -->
<!--     fail("Did you count last names instead of first names?") -->
<!--   } -->
<!--   fail("Not quite. Take a peek at the hint!") -->
<!-- }) -->
<!-- ``` -->

<!-- Now, based on your findings, answer the following question. -->

<!-- ```{r quiz-names} -->
<!-- question("What are the three most common names?", -->
<!--     answer("Michael, Jackson, Jacob"), -->
<!--     answer("Jackson, Jacob, Joseph", -->
<!--            correct = TRUE), -->
<!--     answer("Audrey, Kaitlyn, Joshua"), -->
<!--     answer("Jacob, Jackson, Audrey"), -->
<!--     answer("Kaitlyn, Jacob, Lucas"), -->
<!--     allow_retry = TRUE, -->
<!--     random_answer_order = TRUE -->
<!--   ) -->
<!-- ``` -->

<!-- ### Most common themes -->

<!-- What are the four most common themes of lego sets purchased? Use the code chunk below to find out! -->

<!-- ```{r common-themes, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___(___) -->
<!-- ``` -->

<!-- ```{r common-themes-hint-1} -->
<!-- Look at the previous question for help! -->
<!-- ``` -->

<!-- ```{r common-themes-solution} -->
<!-- lego_sales %>% -->
<!--   count(theme, sort = TRUE) -->
<!-- ``` -->

<!-- ```{r common-themes-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "Star Wars")) { -->
<!--     pass("You have counted themes and sorted the counts correctly.") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Advanced Models ")) { -->
<!--     fail("Did you forget to sort the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Classic")) { -->
<!--     fail("Did you accidentally sort the counts in ascending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Adventure Camp")) { -->
<!--     fail("Did you count subthemes instead of themes?") -->
<!--   } -->
<!--   if(identical(as.numeric(.result[1,2]), 172)) { -->
<!--     fail("Did you count subthemes instead of themes?") -->
<!--   } -->
<!--   fail("Not quite. Take a peek at the hint!") -->
<!-- }) -->
<!-- ``` -->

<!-- Now, based on your findings, answer the following question. -->

<!-- ```{r quiz-themes} -->
<!-- question("What are the four most commonly purchased themes?", -->
<!--     answer("Gear, Duplo, Ninjago, Star Wars"), -->
<!--     answer("Star Wars, Nexo Knights, Gear, City"), -->
<!--     answer("Star Wars, Gear, Mixels, Bionicle"), -->
<!--     answer("Gear, Star Wars, Nexo Knights, Mixels", -->
<!--            correct = TRUE), -->
<!--     answer("Nexo Knights, Gear, Duplo, Friends"), -->
<!--     allow_retry = TRUE, -->
<!--     random_answer_order = TRUE -->
<!--   ) -->
<!-- ``` -->

<!-- ### Most common subtheme -->

<!-- Among the most common theme of lego sets purchased, what is the most common subtheme? -->

<!-- ```{r most-common-subtheme, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___(___) %>% -->
<!--   ___(___) -->
<!-- ``` -->

<!-- ```{r most-common-subtheme-hint-1} -->
<!-- lego_sales %>% -->
<!--   filter(___ == ___) %>% -->
<!--   ___(___) -->
<!-- ``` -->

<!-- ```{r most-common-subtheme-hint-2} -->
<!-- lego_sales %>% -->
<!--   filter(theme == ___) %>% -->
<!--   count(___) -->
<!-- ``` -->

<!-- ``` {r most-common-subtheme-hint-3} -->
<!-- Look at the previous question to see what the most common theme is! -->
<!-- ``` -->

<!-- ```{r most-common-subtheme-solution} -->
<!-- lego_sales %>% -->
<!--   filter(theme == "Star Wars") %>% -->
<!--   count(subtheme, sort = TRUE) -->
<!-- ``` -->

<!-- ```{r most-common-subtheme-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "The Force Awakens")) { -->
<!--     pass("You have successfully counted subthemes and sorted the counts in descending order.") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Battlefront")) { -->
<!--     fail("Did you forget to sort the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Ultimate Collector Series")) { -->
<!--     fail("Did you accidentally sort the counts in ascending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "Star Wars")) {  -->
<!--     fail("Did you count themes instead of subthemes?") -->
<!--   } -->
<!-- fail("Not quite. Take a peek at the hint!") -->
<!-- }) -->
<!-- ``` -->

<!-- Now, based on your findings, answer the following question. -->

<!-- ```{r quiz-subtheme} -->
<!-- question("What is the most common subtheme among the most common theme?", -->
<!--     answer("Buildable Figures"), -->
<!--     answer("Episode V"), -->
<!--     answer("The Force Awakens", -->
<!--            correct = TRUE), -->
<!--     answer("MicroFighters"), -->
<!--     answer("Battlefront"), -->
<!--     allow_retry = TRUE, -->
<!--     random_answer_order = TRUE -->
<!--   ) -->
<!-- ``` -->

<!-- ## Discretizing variables -->

<!-- ### Grouping by age -->

<!-- Create a new variable called `age_group` and group the ages of customers into the following categories: "18 and under", "19 - 25", "26 - 35", "36 - 50", "51 and over". -->

<!-- ``` {r create-age-group, exercise = TRUE} -->
<!-- lego_sales <- ___ %>% -->
<!--   ___(age_group = ___( -->
<!--     ___ -->
<!--   )) -->
<!-- ``` -->

<!-- ``` {r create-age-group-hint-1} -->
<!-- Try using mutate() and case_when()! -->
<!-- ``` -->

<!-- ``` {r create-age-group-hint-2} -->
<!-- Do you need some help with logical operators? Remember -->
<!-- "18 and under" : age <= 18 -->
<!--      "19 - 25" : age >= 19 & age <= 25 -->
<!-- etc. -->
<!-- ``` -->

<!-- ```{r create-age-group-hint-3} -->
<!-- lego_sales <- lego_sales %>% -->
<!--   mutate(age_group = case_when( -->
<!--     age <= 18             ~ "18 and under", -->
<!--     ___ -->
<!--   )) -->
<!-- ``` -->

<!-- ``` {r create-age-group-solution} -->
<!-- lego_sales <- lego_sales %>% -->
<!--   mutate(age_group = case_when( -->
<!--     age <= 18 ~ "18 and under", -->
<!--     age >= 19 & age <= 25 ~ "19 - 25", -->
<!--     age >= 26 & age <= 35 ~ "26 - 35", -->
<!--     age >= 36 & age <= 50 ~ "36 - 50", -->
<!--     age >= 51 ~ "51 and over" -->
<!--   )) -->
<!-- ``` -->

<!-- ```{r create-age-group-check} -->
<!-- grade_this({ -->
<!--   if(identical(paste(as.numeric(table(.result[,"age_group"])), collapse = ", "), "30, 129, 183, 216, 62")) { -->
<!--     pass("You have grouped the ages according to the instructions!") -->
<!--   } -->
<!--   fail("Not quite. Review the hints.") -->
<!-- }) -->
<!-- ``` -->

<!-- ### Most common age group -->

<!-- Count the number of customers in each age group and display the counts in descending order, from most common to least common age group. -->

<!-- Which age group is the most common? Write code in the chunk below to figure it out! -->

<!-- ````{r count-ages, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___(___) -->
<!-- ``` -->
<!-- ````{r count-ages-solution} -->
<!-- lego_sales %>% -->
<!--   count(age_group, sort = TRUE) -->
<!-- ``` -->
<!-- ```{r count-ages-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "36 - 50")) {  -->
<!--     pass("You have successfully determined the most common age group.") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "18 and under")) { -->
<!--     fail("Did you forget to arrange the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "33")) { -->
<!--     fail("Did you count individual ages instead of age groups?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "16")) { -->
<!--     fail("Did you count individual ages instead of age groups?") -->
<!--   } -->
<!--   fail("Not quite. Try using the count function!") -->
<!-- }) -->
<!-- ``` -->
<!-- Now, based on your findings, answer the following question: -->
<!-- ```{r aqe-groups-quiz} -->
<!-- question("Which age group is the most common in the dataset?", -->
<!--     answer("18 and under"), -->
<!--     answer("19 - 25"), -->
<!--     answer("26 - 35"), -->
<!--     answer("36 - 50", -->
<!--            correct = TRUE), -->
<!--     answer("51 and over"), -->
<!--     allow_retry = TRUE, -->
<!--     random_answer_order = TRUE -->
<!--   ) -->
<!-- ``` -->
<!-- ## Grouped summaries  -->
<!-- ### Quantity of purchases by age -->
<!-- Which age group has purchased the highest number of lego sets? -->
<!-- ``` {r sales-by-group, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___(___) %>% -->
<!--   ___( -->
<!--     total_n = ___(___) -->
<!--   ) %>% -->
<!--   arrange(___) -->
<!-- ``` -->
<!-- ``` {r sales-by-group-hint-1} -->
<!-- lego_sales %>% -->
<!--   group_by(___) %>% -->
<!--   summarise( -->
<!--     total_n = sum(___) -->
<!--   ) %>% -->
<!--   arrange(___) -->
<!-- ``` -->
<!-- ``` {r sales-by-group-hint-2} -->
<!-- You need to consider the quantity of purchases. -->
<!-- ``` -->
<!-- ``` {r sales-by-group-hint-3} -->
<!-- lego_sales %>% -->
<!--   group_by(___) %>% -->
<!--   summarise( -->
<!--     total_n = sum(quantity) -->
<!--   ) %>% -->
<!--   arrange(desc(___)) -->
<!-- ``` -->
<!-- ``` {r sales-by-group-solution} -->
<!-- lego_sales %>% -->
<!--   group_by(age_group) %>% -->
<!--   summarise( -->
<!--     total_n = sum(quantity) -->
<!--   ) %>% -->
<!--   arrange(desc(total_n)) -->
<!-- ``` -->
<!-- ```{r sales-by-group-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "36 - 50")) { -->
<!--     pass("Your solution is correct!") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "18 and under")) { -->
<!--     fail("Did you forget to sort the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "33")) { -->
<!--     fail("Did you count individual ages instead of age groups?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "16")) { -->
<!--     fail("Did you count individual ages instead of age groups?") -->
<!--   } -->
<!--   fail("Not quite. Take a peek at the hint!") -->
<!-- }) -->
<!-- ``` -->
<!-- Now, based on your findings, answer the following question: -->
<!-- ```{r quiz-sales-by-group} -->
<!-- question("Which age group has purchased the highest number of lego sets?", -->
<!--     answer("18 and under"), -->
<!--     answer("19 - 25"), -->
<!--     answer("26 - 35"), -->
<!--     answer("36 - 50", -->
<!--            correct = TRUE), -->
<!--     answer("51 and over"), -->
<!--     allow_retry = TRUE -->
<!--   ) -->
<!-- ``` -->
<!-- ### Most generous age group -->
<!-- Which age group has spent the most money on legos? -->
<!-- ```{r age-group-money, exercise = TRUE} -->
<!-- lego_sales %>% -->
<!--   ___( -->
<!--     amount_spent = ___ * ___ -->
<!--   ) %>% -->
<!--   ___(___) %>% -->
<!--   ___( -->
<!--     total_spent = ___(___) -->
<!--   ) %>% -->
<!--   arrange(___) -->
<!-- ``` -->
<!-- ````{r age-group-money-hint-1} -->
<!-- You will need to consider quantity of purchases as well as price of lego sets. -->
<!-- ``` -->
<!-- ```{r age-group-money-hint-2} -->
<!-- lego_sales %>% -->
<!--   mutate( -->
<!--     amount_spent = ___ * ___ -->
<!--   ) %>% -->
<!--   group_by(___) %>% -->
<!--   summarise( -->
<!--     total_spent = ___(___) -->
<!--   ) %>% -->
<!--   arrange(desc(___)) -->
<!-- ``` -->
<!-- ```{r age-group-money-hint-3} -->
<!-- lego_sales %>% -->
<!--   mutate( -->
<!--     amount_spent = us_price * quantity -->
<!--   ) %>% -->
<!--   group_by(age_group) %>% -->
<!--   summarise( -->
<!--     total_spent = sum(___) -->
<!--   ) %>% -->
<!--   arrange(desc(___)) -->
<!-- ``` -->
<!-- ```{r age-group-money-solution} -->
<!-- lego_sales %>% -->
<!--   mutate( -->
<!--     amount_spent = us_price * quantity -->
<!--   ) %>% -->
<!--   group_by(age_group) %>% -->
<!--   summarise( -->
<!--     total_spent = sum(amount_spent) -->
<!--   ) %>% -->
<!--   arrange(desc(total_spent)) -->
<!-- ``` -->
<!-- ```{r age-group-money-check} -->
<!-- grade_this({ -->
<!--   if(identical(as.character(.result[1,1]), "36 - 50")) { -->
<!--     pass("Your solution is correct!") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "18 and under")) { -->
<!--     fail("Did you forget to sort the counts in descending order?") -->
<!--   } -->
<!--   if(identical(as.character(.result[1,1]), "57")){ -->
<!--     fail("Did you count individual ages instead of age groups?") -->
<!--   } -->
<!--   fail("Not quite. Take a peek at the hint!") -->
<!-- }) -->
<!-- ``` -->
<!-- Now, based on your findings, answer the following question: -->
<!-- ```{r spend-by-group} -->
<!-- question("Which age group has spent the most money on lego sets?", -->
<!--     answer("18 and under"), -->
<!--     answer("19 - 25"), -->
<!--     answer("26 - 35"), -->
<!--     answer("36 - 50", -->
<!--            correct = TRUE), -->
<!--     answer("51 and over"), -->
<!--     allow_retry = TRUE, -->
<!--     random_answer_order = TRUE -->
<!--   ) -->
<!-- ``` -->
<!-- ## Wrap up -->
<!-- Congratulations, you have completed tutorial 4!  -->
<!-- We hope you enjoyed reviewing what we've learned so far! -->